## Sprint Challenge: Data Wrangling and Storytelling
### Notebook points total: 14


## Python Fundamentals

**Task 1** - Python Objects
* Create a list object called `list_practice` using the following three strings: `bloom`, `data`, `python` 

<!-- grade_id: cell-6ee0685279505899, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
list_practice= ['bloom', 'data', 'python']
```

**Task 1 Test**

<!-- grade_id: cell-7f5ed8ccc5b15f71, points: 1, grade: true, solution: false -->
```python
#Task 1 - Test
assert isinstance(list_practice, list), "Make sure you created a list object"

```


**Task 2** - Dictionaries
* Create a dictionary object called `diction_practice`. 
* Assign the following values to their respective keys listed above: `tech`,  `science`, `language`

*Hint:* There are multiple ways you can accomplish this task. You can either write out the key:values pairs manually, or iteratively combine two lists using the [`zip` function](https://docs.python.org/3/library/functions.html).

<!-- grade_id: cell-d2aac83e74b8fb89, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
temp_list = ['tech','science','language']
zip_iter = zip(list_practice,temp_list)
diction_practice = dict(zip_iter)
print(diction_practice)
```

Output:
```
{'bloom': 'tech', 'data': 'science', 'python': 'language'}
```

**Task 2 - Test**

<!-- grade_id: cell-6f3011fa357d270f, points: 1, grade: true, solution: false -->
```python
#Task 2 - Test

assert isinstance(diction_practice, dict), "Did you use the correct syntax?"

```


**Task 3** - Dictionaries
* Reassign the value of the `python` key in your dictionary to `'programming_language'`

<!-- grade_id: cell-59c6653a1a1faae6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
diction_practice['python'] = 'programming_language'
diction_practice
```

Output:
```
{'bloom': 'tech', 'data': 'science', 'python': 'programming_language'}
```

**Task 3 Tests**

<!-- grade_id: cell-94bd4331d6cb29c9, points: 1, grade: true, solution: false -->
```python
#Task 3 Tests

assert diction_practice['data'] == 'science', "Make sure your values are assigned to their correct keys"

```

## Use the following information to complete Tasks 



In this Sprint Challenge you will first "wrangle" some data from [Gapminder](https://www.gapminder.org/about-gapminder/), a Swedish non-profit co-founded by Hans Rosling. "Gapminder produces free teaching resources making the world understandable based on reliable statistics."
- [Cell phones (total), by country and year](https://raw.githubusercontent.com/open-numbers/ddf--gapminder--systema_globalis/master/countries-etc-datapoints/ddf--datapoints--cell_phones_total--by--geo--time.csv)
- [Population (total), by country and year](https://raw.githubusercontent.com/open-numbers/ddf--gapminder--systema_globalis/master/countries-etc-datapoints/ddf--datapoints--population_total--by--geo--time.csv)
- [Geo country codes](https://github.com/open-numbers/ddf--gapminder--systema_globalis/blob/master/ddf--entities--geo--country.csv)

These two links have everything you need to successfully complete the first part of this sprint challenge.
- [Pandas documentation: Working with Text Data](https://pandas.pydata.org/pandas-docs/stable/text.html) (one question)
- [Pandas Cheat Sheet](https://github.com/pandas-dev/pandas/blob/master/doc/cheatsheet/Pandas_Cheat_Sheet.pdf) (everything else)


**Task 4** - Load and print the cell phone data.  Pandas and numpy import statements have been included for you.

* Load your CSV file found at `cell_phones_url` into a DataFrame object named `cell_phones`


<!-- grade_id: cell-6f16afb9d271f949, points: 0, grade: false, solution: true -->
```python
# Task 4

# Imports 
import pandas as pd
import numpy as np

cell_phones_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Cell__Phones/cell_phones.csv'

# Load the dataframe and print the top 5 rows

# YOUR CODE HERE
cell_phones = pd.read_csv(cell_phones_url)

```

**Task 4 Test**

<!-- grade_id: cell-226f7bf8e9ea24f9, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert isinstance(cell_phones, pd.DataFrame), 'Have you created a DataFrame named `cell_phones`?'
assert len(cell_phones) == 9574

```

**Task 5** - Load and print the population data.  

* Load the CSV file found at `population_url` into a DataFrame named `population`



<!-- grade_id: cell-561c2d59728188a9, points: 0, grade: false, solution: true -->
```python
# Task 5

population_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Population/population.csv'

# Load the dataframe and print the first 5 records

# YOUR CODE HERE
population = pd.read_csv(population_url)
```

**Task 5 Test**

<!-- grade_id: cell-59d01cd695becd74, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

assert isinstance(population, pd.DataFrame), 'Have you created a DataFrame named `population`?'
assert len(population) == 59297

```

**Task 6** - Load and print the geo country codes data.  

* Load the CSV file found at `geo_codes_url` into a DataFrame named `geo_codes`


<!-- grade_id: cell-eb4d290384535503, points: 0, grade: false, solution: true -->
```python
# Task 6

geo_codes_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/GEO_codes/geo_country_codes.csv'

# Load the dataframe and print out the first 5 records

# YOUR CODE HERE
geo_codes = pd.read_csv(geo_codes_url)
```

**Task 6 Test**

<!-- grade_id: cell-39240405659c0c19, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

assert geo_codes is not None, 'Have you created a DataFrame named `geo_codes`?'
assert len(geo_codes) == 273

```

**Task 7** - Check for missing values

Let's check for missing values in each of these DataFrames: `cell_phones`, `population` and `geo_codes`

* Check for missing values in the following DataFrames:
    * Assign the total number of missing values in `cell_phones` to the variable `cell_phones_missing`
    * Assign the total number of missing values in `population` to the variable `population_missing`
    * Assign the total number of missing values in `geo_codes` to the variable `geo_codes_missing` 
        * Hint: you will need to do a sum of a sum for this last task.

<!-- grade_id: cell-9426cd5765574e07, points: 0, grade: false, solution: true -->
```python
# Task 7

# Check for missing data in each of the DataFrames

# YOUR CODE HERE
cell_phones_missing = cell_phones.isnull().sum().sum()
population_missing = population.isnull().sum().sum()
geo_codes_missing = geo_codes.isnull().sum().sum()

```

Output:
```
781
```

**Task 7 Test**

<!-- grade_id: cell-cf6ab3b4b1e8afc1, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

if geo_codes_missing == 21: print('ERROR: Make sure to use a sum of a sum for the missing geo codes!')  

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Merge the `cell_phones` and `population` DataFrames.

* Merge the `cell_phones` and `population` dataframes with an **inner** merge on both the `geo` and `time` columns.
* Call the resulting dataframe `cell_phone_population`

<!-- grade_id: cell-decaebaa844aa3a5, points: 0, grade: false, solution: true -->
```python
# Task 8

# Merge the cell_phones and population dataframes

# YOUR CODE HERE
cell_phone_population = cell_phones.merge(population, on=['geo','time'])

```

**Task 8 Test**

<!-- grade_id: cell-dd2473ea91f15f30, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

assert cell_phone_population is not None, 'Have you merged created a DataFrame named cell_phone_population?'
assert len(cell_phone_population) == 8930

```

**Task 9** - Merge the `cell_phone_population` and `geo_codes` DataFrames

* Merge the `cell_phone_population` and `geo_codes` DataFrames with an inner merge using the `geo` column.
* **Only merge the `country` and `geo` columns from the `geo_codes` dataframe.** 
* Call the resulting DataFrame `geo_cell_phone_population`


<!-- grade_id: cell-1ce5a2360ee6fd20, points: 0, grade: false, solution: true -->
```python
# Task 9

# Merge the cell_phone_population and geo_codes dataframes
# Only include the country and geo columns from geo_codes

# YOUR CODE HERE
geo_cell_phone_population = cell_phone_population.merge(geo_codes[['country','geo']], on='geo')
geo_cell_phone_population
```

Output:
```
geo  time  cell_phones_total  population_total      country
0     afg  1960                0.0           8996967  Afghanistan
1     afg  1965                0.0           9956318  Afghanistan
2     afg  1970                0.0          11173654  Afghanistan
3     afg  1975                0.0          12689164  Afghanistan
4     afg  1976                0.0          12943093  Afghanistan
...   ...   ...                ...               ...          ...
8925  zwe  2015         12757410.0          13814642     Zimbabwe
8926  zwe  2016         12878926.0          14030338     Zimbabwe
8927  zwe  2017         14092104.0          14236599     Zimbabwe
8928  zwe  2018         12908992.0          14438812     Zimbabwe
8929  zwe  2019         13195902.0          14645473     Zimbabwe

[8930 rows x 5 columns]
```

**Task 9 Test**

<!-- grade_id: cell-764d5b72ae382339, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test
assert len(geo_cell_phone_population) == 8930
assert type(geo_cell_phone_population) == pd.DataFrame

```

**Task 10** - Calculate the number of cell phones per person.

* Use the `cell_phones_total` and `population_total` columns to calculate the number of cell phones per person for every year. (In other words, for every row). 
* Call this new feature (column) `phones_per_person` and add it to the `geo_cell_phone_population` DataFrame (you'll be adding the column to the DataFrame).


<!-- grade_id: cell-bd0bce2920604643, points: 0, grade: false, solution: true -->
```python
# Task 10

# YOUR CODE HERE
geo_cell_phone_population['phones_per_person']= geo_cell_phone_population['cell_phones_total'] / geo_cell_phone_population['population_total']
geo_cell_phone_population
```

Output:
```
geo  time  ...      country  phones_per_person
0     afg  1960  ...  Afghanistan           0.000000
1     afg  1965  ...  Afghanistan           0.000000
2     afg  1970  ...  Afghanistan           0.000000
3     afg  1975  ...  Afghanistan           0.000000
4     afg  1976  ...  Afghanistan           0.000000
...   ...   ...  ...          ...                ...
8925  zwe  2015  ...     Zimbabwe           0.923470
8926  zwe  2016  ...     Zimbabwe           0.917934
8927  zwe  2017  ...     Zimbabwe           0.989850
8928  zwe  2018  ...     Zimbabwe           0.894048
8929  zwe  2019  ...     Zimbabwe           0.901023

[8930 rows x 6 columns]
```

**Task 10 Test**

<!-- grade_id: cell-45c955c43e471400, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 11** - Identify the number of cell phones per person in the US in 2017

* Create a one-row subset of `geo_cell_phone_population` with data on cell phone ownership in the United States for the year 2017.
* Call this subset DataFrame `US_2017`.
* Print `US_2017`.

<!-- grade_id: cell-665e83d11e594d90, points: 0, grade: false, solution: true -->
```python
# Task 11

# Determine the number of cell phones per person in the US in 2017

# YOUR CODE HERE
tempframe = geo_cell_phone_population.loc[geo_cell_phone_population['geo']=='usa']
US_2017 = tempframe.loc[tempframe['time']==2017]
US_2017
# View the DataFrame

```

Output:
```
geo  time  ...        country  phones_per_person
8455  usa  2017  ...  United States           1.230448

[1 rows x 6 columns]
```

**Task 11 Test**

<!-- grade_id: cell-ea08fdda80cf9731, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 12** - Describe the numeric variables in `geo_cell_phone_population`

* Calculate the summary statistics for the quantitative variables in `geo_cell_phone_population` using `.describe()`.
* Find the mean value for `phones_per_person` and assign it to the variable `mean_phones`. Define your value out to two decimal points.


<!-- grade_id: cell-181c9805c52dfda8, points: 0, grade: false, solution: true -->
```python
# Task 12

# YOUR CODE HERE
geo_cell_phone_population.describe()
mean_phones = geo_cell_phone_population['phones_per_person'].mean()

```

**Task 12 Test**

<!-- grade_id: cell-0a0cfb1ac7acc279, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 13** - Subset the DataFrame for 2017

* Create a new dataframe called `df2017` that includes **only** records from `geo_cell_phone_population` that ocurred in 2017.

<!-- grade_id: cell-f3fa17e15b5174ac, points: 0, grade: false, solution: true -->
```python
# Task 13

# Create a new dataframe called df2017 that includes only records from geo_cell_phone_population that ocurred in 2017.

# YOUR CODE HERE
df2017 = geo_cell_phone_population[geo_cell_phone_population['time'] == 2017]
df2017
```

Output:
```
geo  time  ...               country  phones_per_person
45    afg  2017  ...           Afghanistan           0.659291
93    ago  2017  ...                Angola           0.446861
141   alb  2017  ...               Albania           1.257104
189   and  2017  ...               Andorra           1.043378
227   are  2017  ...  United Arab Emirates           2.089785
...   ...   ...  ...                   ...                ...
8738  wsm  2017  ...                 Samoa           0.635812
8784  yem  2017  ...                 Yemen           0.551750
8831  zaf  2017  ...          South Africa           1.552324
8879  zmb  2017  ...                Zambia           0.797369
8927  zwe  2017  ...              Zimbabwe           0.989850

[194 rows x 6 columns]
```

**Task 13 Test**

<!-- grade_id: cell-b19a6956b5dddadb, points: 1, grade: true, solution: false -->
```python
# Task 13 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 14** - Identify the five countries with the most cell phones per person in 2017

* Sort the `df2017` DataFrame by `phones_per_person` in descending order and assign the result to `df2017_top`. Your new DataFrame should only have **five** rows (Hint: use `.head()` to return only five rows).
* Print the first 5 records of `df2017_top`.

<!-- grade_id: cell-c9b701bfe897edf5, points: 0, grade: false, solution: true -->
```python
# Task 14

# Sort the df2017 dataframe by phones_per_person in descending order
# Return only five (5) rows

# YOUR CODE HERE
df2017_top = df2017.sort_values('phones_per_person', ascending=False).head()

# View the df2017_top DataFrame
df2017_top
```

Output:
```
geo  time  ...               country  phones_per_person
3448  hkg  2017  ...      Hong Kong, China           2.517652
227   are  2017  ...  United Arab Emirates           2.089785
365   atg  2017  ...   Antigua and Barbuda           1.928216
5253  mdv  2017  ...              Maldives           1.813303
1937  cri  2017  ...            Costa Rica           1.785944

[5 rows x 6 columns]
```

**Task 14 Test**

<!-- grade_id: cell-5f6fffc9db1b9492, points: 1, grade: true, solution: false -->
```python
# Task 14 - Test

assert df2017_top.shape == (5,6), 'Make sure you return only five rows'

```

```python

```